<a href="https://colab.research.google.com/github/YKannikanti/Yasasri-/blob/main/week06/week06_assignment.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 30.1 List of Allergies

In the /data directory on the Jupyter server, there is a file called `allergies.json` that contains a list of patient allergies.  It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter by clicking the _Home_ icon, then the _from_instructor_ folder, and then the _data_ folder.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

```python
import json
ALLERGIES_FILE= "allergies.json"
```

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
def patient_allergy_count(json_file):
    with open(json_file) as f:
      data = json.load(f)
      if 'entry' in data:
        return len(data['entry'])
      return 0
```

```python
allergy_count = patient_allergy_count(ALLERGIES_FILE)
```

```python
assert type(allergy_count) == int
assert allergy_count == 4
```

### 30.2 Number of Patients

If you dig a little bit deaper into this list of allergies, you'll see that each result has a patient associated with it.  Create a funcation called `patient_count(json_file)` that will count how many unique patients we have in this JSON structure.  

```python
import json
ALLERGIES_FILE= "allergies.json"
```

```python
def patient_count(json_file):
  with open(json_file) as f:
    data = json.load(f)
    patients = set()
    for entry in data['entry']:
      patients.add(entry['resource']['patient']['display'])
  return len(patients)

```

```python
assert type(patient_count(ALLERGIES_FILE)) == int
assert patient_count(ALLERGIES_FILE) == 2
```

### 30.3 How Many Allergies per Patient

Although each entry is a separate allergy, several of them are for the same patient.  Write a function called `allergy_per_patient(json_file)` that counts up how many allergies each patient has.


```python
import json
ALLERGIES_FILE= "allergies.json"
```

```python
def allergy_per_patient(json_file):
  with open(json_file) as f:
    data = json.load(f)
    allergies_per_patient = {}
    for entry in data['entry']:
      patient = entry['resource']['patient']['display']
      if patient not in allergies_per_patient:
        allergies_per_patient[patient] = 0
      allergies_per_patient[patient] += 1
  return allergies_per_patient
```

```python
assert type(allergy_per_patient(ALLERGIES_FILE)) == dict
assert allergy_per_patient(ALLERGIES_FILE) == {'Paul Boal': 1, 'Jason Argonaut': 3}
```

### 30.4 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
import json
```

```python
def allergy_list(json_file):
  with open(json_file) as f:
    data = json.load(f)
    allergies = []
    for entry in data['entry']:
      patient = entry['resource']['patient']['display']
      substance = entry['resource']['substance']['text']
      if 'reaction' in entry['resource']:
        reaction = entry['resource']['reaction'][0]['manifestation'][0]['text']
      else:
         reaction = None
      allergies.append([patient, substance, reaction])
  return sorted(allergies)

```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
  assert allergy_list(ALLERGIES_FILE) == [
 ['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']
 ]

```

### 30.5 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  You can solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
import json

ALLERGIES_FILE= "allergies.json"
```

```python
def allergy_list(json_file):
  with open(json_file) as f:
    data = json.load(f)
    allergies = []
    for entry in data['entry']:
      patient = entry['resource']['patient']['display']
      substance = entry['resource']['substance']['text']
      if 'reaction' in entry['resource']:
        reaction = entry['resource']['reaction'][0]['manifestation'][0]['text']

      else:
         reaction = None
      allergies.append([patient, substance, reaction])
  return sorted(allergies)

def allergy_reaction(json_file, patient, substance):
  allergies = allergy_list(json_file)
  for allergy in allergies:
    if allergy[0] == patient and allergy[1] == substance:
      return allergy[2]
  return None

```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') ==  'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week06/week06_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

